Hypothalamic representation of the imminence of predator threat detected by the vomeronasal organ in mice

Author response:

The following is the authors’ response to the previous reviews

We greatly appreciate all the reviewers’ constructive comments on our previously revised manuscript. In the current revision, we added several experimental data for answering the reviewers’ comments. Below we describe our point-by-point responses to their comments:

Reviewer #1 (Public Review):

Unaddressed and additional concerns (re-submission)

In this revised version of the manuscript, the authors have made important modifications in the text, inserted new references, and incorporated additional quantifications of cFos immunolabeling in three brain regions, as recommended by the reviewers. While these modifications have significantly improved the quality of the manuscript, other critical concerns raised during the initial submission of the

manuscript (Major concerns 1, 2, and 4; some of them also raised by the other reviewers) were not properly addressed by the authors. On several occasions, the authors recognize the importance of clarifying the points for the correct interpretation of the results but opt for leaving the open questions to be addressed during future studies. Therefore, the authors might consider adding a new section at the end of the manuscript to include all the caveats and future directions.

In the current revision, in order to answer the reviewer #1’s original concerns 1, 2, and 4, we added several experimental data.

Original major concerns 1) and 2): Regarding whether mice are detecting qualitative or quantitative differences between fresh and old cat saliva.

To address these concerns, as shown in new Figure 1I and J, we measured volumes of saliva contained in in individual swabs and total protein concentrations at the time of behavior tests: Fresh (15 minutes after collection) and Old (4 hours after collection). The saliva volumes at the time of behavioral testing were indistinguishable between fresh and old samples (Figure 1I). In addition, the concentrations of total proteins in both fresh and old saliva were also indiscernible (Figure 1J). Furthermore, we also examined the difference of the amount of Fel d 4 protein, one of the most abundant proteins in cat saliva, between fresh and old saliva by conducting western blotting analyses. As shown in new Supplemental Figure 2, the amount of Fel d 4 was nearly equivalent between fresh and old saliva. Indeed, our analyses using recombinant Fel d 4 protein showed that Fel d 4 does not induce freezing behavior (Supplemental Figure 5). Based on these findings, we believe that the difference between fresh and old cat saliva lies in specific components rather than the total or major saliva content. One possible explanation for this difference is the time-dependent reduction of specific freezing-inducing components in old saliva.

To investigate such a possibility, we also examined mouse behavior directed toward swabs containing diluted fresh cat saliva. Indeed, exposure to diluted fresh saliva resulted in a shorter duration of freezing behavior. Fresh saliva diluted to 70% induced freezing behavior for a duration equivalent to that of undiluted fresh saliva, while freezing behavior in response to 50% and 30% fresh saliva was significantly reduced to the same duration as that observed with old saliva (Figure 1K). The duration of direct interaction with swabs containing 70% and 50–30% fresh saliva also exhibited a similar trend to that observed with fresh and old saliva swabs, respectively (Figure 1L).

These new results provide compelling evidence that the differential freezing response of mice to fresh versus old cat saliva is not attributed to quantitative differences, such as total volume, total protein concentration, or the amount of major proteins like Fel d 4. However, when fresh saliva was diluted, we observed a corresponding reduction in freezing behavior, suggesting that specific components within the saliva—those responsible for inducing freezing—may decrease over time.

Our findings indicate that while the overall content of saliva remains consistent over time, specific freezing-inducing components seem to degrade or reduce at a different rate than other components, which alters the composition of saliva over time. The speed of reduction of these freezing-inducing components appears to be different from more stable proteins such as Fel d 4. As a result, the composition of saliva changes over time, leading to a qualitative difference between fresh and old saliva that mice can detect. This ability to discern such subtle chemical changes likely reflects an adaptive sensory mechanism, allowing mice to respond to predator cues to induce optimal defensive behavior in a certain context. Identifying the specific freezing-inducing components through traditional purification processes, such as high-performance liquid chromatography followed by behavioral examination (Haga-Yamanaka et al., 2014; Kimoto et al., 2005), is crucial for a deeper understanding of the mechanisms underlying the observed behavior. Our research team is actively working to isolate these molecules, and we hope to report our findings in future studies.

(4) The interpretation that fresh and old saliva activates different subpopulations of neurons in the VMH based on the observation that cFos positively correlates with freezing responses only with the fresh saliva lacks empirical evidence. To address this question, the authors should use two neuronal activity markers to track the response of the same population of VHM cells within the same animals during exposure to fresh vs. old saliva.

To address this issue, as shown in the new Figure 7, we performed a double exposure experiment using Fos2A-iCreERT2; Ai9 (TRAP2) mice (Allen et al., 2017; DeNardo et al., 2019). In this experiment, mice were exposed to the first stimulus under the treatment of 4-hydroxytamoxifen (4-OHT). One week after the initial exposure, the same mice were subjected to a second stimulus exposure for one hour. Through this paradigm, neurons activated by the first stimulus were visualized by tdTomato, while ones activated by the second stimulus were detected as cFos-IR (Figure 7A). Quantification of tdTomato and cFos-IR double-positive cells among tdTomato-labeled cells revealed that 43% (mean per animal: 61 / 143) of cells activated by fresh saliva during the first exposure were also activated by fresh saliva during the second exposure, whereas only 16% (17 / 106) of cells activated by old saliva during the first exposure were activated by fresh saliva during the second exposure (p = 7.5e-6, Chi-squared test). The difference in the fraction of overlapping cells between fresh and old saliva exposures was found significant when we compared the two groups of animals (Figure 7D, p = 0.0035, permutation test). Additionally, quantification of tdTomato and cFos-IR double-positive cells among cFos-IR cells indicated that over 27% (61 / 226) of cells activated by fresh saliva during the second exposure were previously activated by fresh saliva, whereas only 15% (17 / 112) of cells activated by fresh saliva during the second exposure were previously activated by old saliva (p = 0.015, Chi-squared test). The difference in the fraction of overlapping cells between fresh and old saliva exposures was also significant in this analysis (Figure 7E,p = 0.0060, permutation test). Together, these results demonstrate that fresh and old cat saliva activate largely different populations of neurons within the VMH. These new results were described on page 11 line 18 – page 12 line 8.

In addition to these unaddressed concerns, some new issues have emerged in the new version of the manuscript. For example, the following paragraph introduced in the discussion section is not supported by the experimental findings.

"We assume that such differential activations of the mitral cells between fresh and old saliva result in the differential activation of targeting neural substrates, possibly MeApv, which results in differential activation of VMH neurons (Figure 7)."

Although the authors did not observe statistical differences in cFos expression in the pvMeA among groups, they claim that the differences in cFos expression in the VMH between fresh vs. old saliva are mediated by differential activation of upstream neurons in the MeApv. The lack of statistical differences may be caused by the reduced number of subjects in each group, as recognized in the text by the

authors.

We appreciate the reviewer's thoughtful comment. We agree that the paragraph in the comment, which presented a working hypothesis regarding differential activations of mitral cells and the MeApv between fresh and old saliva exposures, was speculative and not fully supported by our experimental findings. To address this, we have removed the assumptions related to the differential responses of mitral cells and the MeApv from the discussion and have updated the figure accordingly (now presented as new Figure 8).

Moreover, the authors propose that in addition to fel d 4, multiple molecules present in the cat saliva can be inducing distinct defensive responses in the animals, but they do not provide any reference to support their claim.

We thank the reviewer for highlighting this point. Our claim regarding the presence of other molecules in cat saliva inducing freezing defensive responses is based on our observation, as shown in the new Supplemental Figure 5, that recombinant Fel d 4 protein alone does not induce freezing behavior. This suggests the existence of other unidentified components in cat saliva that may contribute to freezing behavior. As we agree that identifying these specific freezing-inducing components is important for a more comprehensive understanding of the underlying mechanisms, our research team is actively working to isolate these molecules, and we hope to report our findings in future studies.

Reviewer #2 (Public Review):

The findings are relatively preliminary. The identities of the receptor and the ligand in the cat saliva that induces the behavior remain unclear. The identity of VMH cells that are activated by the cat saliva remains unclear. There is a lack of targeted functional manipulation to demonstrate the role of V2R-A4 or VMH cells in the behavioral response to the cat saliva.

We thank the reviewer’s important insight on the need for further investigation into the molecular and neural mechanisms underlying the behavioral response to cat saliva. We recognize the importance of conducting studies involving V2R-A4 receptor knockouts and targeted functional manipulations within the VMH using neural circuit perturbation approaches.

However, the V2R-A4 subfamily consists of 25 Vmn2r genes, most of which are closely grouped together, forming a V2R-A4 gene cluster within a 2.5-megabase chromosomal region. As we described in our recent review article (Rocha et al., 2024), the Vmn2r genes within the V2R-A4 subfamily display a high degree of homology, with nucleotide and amino acid identities among the several Vmn2rs surpassing 97-99%, suggesting possible redundancy among these receptor genes. This is in stark contrast to the diversity typically observed within other V2R subfamilies. Consequently, knockout strategies targeting a single receptor gene, which have been successful for other vomeronasal receptors, may not be effective for V2R-A4 receptor genes. The most appropriate strategy for examining the necessity of V2R-A4 receptors would be knocking out the entire V2R-A4 gene cluster, spanning a 2.5-megabase chromosomal region. Due to the technical challenges involved, addressing this issue is not feasible in the foreseeable future. Moreover, in our current study, we aimed to establish the foundational relationship between predator cues in cat saliva and defensive behaviors. We view our findings as an important first step that sets the stage for these more targeted and mechanistic studies involving the neural circuit perturbation experiments, such as optogenetics and Designer Receptors Exclusively Activated by Designer Drugs (DREADDs), in the next step.

Reviewer #3 (Public Review):

Weaknesses:

(1) It is unclear if fresh and old saliva indeed alter the perceived imminence of predation, as claimed by the authors. Prior work indicates that lower imminence induces anxiety-related actions, such as re- organization of meal patterns and avoidance of open spaces, while slightly higher imminence produces freezing. Here, the authors show that fresh and old predator saliva only provoke different amounts of freezing, rather than changing the topography of defensive behaviors, as explained above. Another prediction of predatory imminence theory would be that lower imminence induced by old saliva should produce stronger cortical activation, while fresh saliva would activate amygdala, if these stimuli indeed correspond to significantly different levels of predation imminence.

We appreciate the reviewer’s insightful comments regarding the perceived imminence of predation and the behavioral responses to fresh and old saliva. Our study specifically focused on comparing the defensive behaviors of mice in response to 15-minute-old and 4-hour-old cat saliva, particularly within the context of freezing behavior in their home cages. We chose these specific time points to capture the potential variation in behavioral intensity rather than the full spectrum of defensive behaviors. While a more comprehensive analysis—including varying time points, different types of defensive behaviors, and broader neural activation patterns (e.g., cortical versus amygdala activation)—might provide further insights into predation imminence theory, these aspects were beyond the scope of our current study. Future research could certainly address these points by examining behavioral and neural responses across additional saliva aging intervals and in varied behavioral contexts. Such studies would complement and extend the findings presented here, further elucidating the relationship between predator cue characteristics and defensive behaviors.

(2) It is known that predator odors activate and require AOB, VNO and VMH, thus replications of these findings are not novel, decreasing the impact of this work.

As the reviewer mentioned, the activation of the AOB, VNO, and VMH by predator odors has been established in prior studies. However, our study provides new insights by demonstrating that defensive freezing behavior in response to predator odors is mediated through the vomeronasal organ (VNO) sensory circuit, which has not been previously shown. The novelty of our work lies in two key findings: 1) the introduction of a new behavioral paradigm that assesses freezing responses to predator cues based on the freshness of chemosensory signals in cat saliva, and 2) the demonstration that the vomeronasal sensory circuit mediates defensive freezing behavior in response to cat saliva.

Additionally, our results show that cat saliva of different freshness levels differentially activates VNO sensory neurons that express the same subfamily of sensory receptors. This differential activation subsequently modulates the downstream neural circuits, leading to varied freezing behavioral outcomes. We believe these findings provide a novel conceptual advance over previous studies by elucidating a more detailed mechanism of how predator-derived cues influence defensive behaviors through the accessory olfactory system.

(3) There is a lack of standard circuit dissection methods, such as characterizing the behavioral effects of increasing and decreasing neural activity of relevant cell bodies and axonal projections, significantly decreasing the mechanistic insights generated by this work

We thank the reviewer for this valuable comment. Investigating the behavioral effects of manipulating specific cell types and axonal projections, as well as characterizing circuit connectivity, is essential for a more comprehensive understanding of the underlying neural circuits. These approaches, such as modulating neural activity in defined cell populations and dissecting circuit pathways, using optogenetics, DREADD, etc., would provide deeper mechanistic insights. In our current study, however, we aimed to establish the foundational relationship between predator cues in cat saliva and defensive behaviors. We view our findings as an important first step that sets the stage for these more targeted and mechanistic studies in the future.

(4) The correlation shown in Figure 5c may be spurious. It appears that the correlation is primarily driven by a single point (the green square point near the bottom left corner). All correlations should be calculated using Spearman correlation, which is non-parametric and less likely to show a large correlation due to a small number of outliers. Regardless of the correlation method used, there are too few points in Figure 5c to establish a reliable correlation. Please add more points to 5c.

We appreciate the reviewer’s suggestion regarding the correlation analysis in Figure 5E. We assessed the normality of our data using both the Shapiro-Wilk and Kolmogorov-Smirnov tests, which confirmed that the dataset is parametric, justifying the use of a parametric correlation method in this context. However, we acknowledge the concern about the limited number of data points and the influence of potential outliers on the observed correlation. Increasing the sample size might provide a more robust assessment of correlation patterns and reduce the potential impact of any single data point. While this would be an important direction for future research, such as with larger sample sizes, it is beyond the scope of the current study.

(5) Please cite recent relevant papers showing VMH activity induced by predators, such as https://pubmed.ncbi.nlm.nih.gov/33115925/ and https://pubmed.ncbi.nlm.nih.gov/36788059/

We thank the reviewer’s suggestion to cite these important papers. https://pubmed.ncbi.nlm.nih.gov/33115925/ (Esteban Masferrer et al., 2020) and https://pubmed.ncbi.nlm.nih.gov/36788059/ (Tobias et al., 2023) are now cited at page 16 line 10 in Discussion under “Differential activation of VMH neurons potentially underlying distinct intensities of freezing behavior.”

(6) Add complete statistical information in the figure legends of all figures, which should include n, name of test used and exact p values.

We included statistical analysis results in figure legends; for Figure 6B, we provided statistical analysis results in Supplemental Table 1.

(7) Some of the findings are disconnected from the story. For example, the authors show V2R-A4- expressing cells are activated by predator odors. Are these cells more likely to be connected to the rest of the predatory defense circuit than other VNO cells?

Yes, our hypothesis posits that V2R-A4-expressing VNO sensory neurons serve as receptor neurons for predator cues present in cat saliva. Additionally, we assume that these specific sensory neurons have stronger anatomical connections with the defensive circuit compared to VNO sensory neurons expressing other receptor subfamilies. In our modified Discussion section, we discussed this point under “V2R-A4 subfamily as the receptor for predator cues in cat saliva.”

(8) Please paste all figure legends directly below their corresponding figure to make the manuscript easier to read

We have added figure legends directly below their corresponding figures.

(9) Were there other behavioral differences induced by fresh compared to old saliva? Do they provoke differences in stretch-attend risk evaluation postures, number of approaches, average distance to odor stimulus, velocity of movements towards and away the odor stimulus, etc?

We appreciate the reviewer's valuable comments. We have now incorporated an analysis of stretch-sniff risk assessment behavior, presented in new Figure 1F (graph) and Supplemental Figure 1B (raster plot). Mice exhibited stretch-sniff risk assessment behavior, which remained consistent across control, fresh saliva, and old saliva swabs. Additionally, we have also included a raster plot for direct investigation, previously noted as ‘interaction’ in the original manuscript (Supplemental Figure 1C). Mice exposed to a swab containing either fresh or old saliva significantly avoided directly investigating the swab. In contrast, mice exposed to a clean control swab spent a significant amount of time directly investigating the swab, engaging in behaviors such as sniffing and chewing (Figure 1G). A comparison of temporal behavioral patterns revealed a slightly higher frequency of direct investigation behavior toward old saliva compared to fresh saliva at the beginning of the exposure period (Supplemental Figure 1C).

Reviewer #3 (Recommendations For The Authors):

The authors have partially addressed several important points raised in the prior review, increasing the strength of the manuscript. However, 2 key questions already raised previously, were not addressed:

(1) Is old saliva qualitatively different from new saliva, or is it the same as a smaller amount of new saliva? As Reviewer 1 wrote: "An important point that the authors should clarify in this study is whether mice are detecting qualitative or quantitative differences between fresh and old cat saliva."

Since one of the author's main points is that fresh and old saliva elicit different perceived threat imminences, it is crucial to show that these two stimuli are somehow qualitatively different.

One way to investigate this could be to show that animals perform different behaviors when exposed to smaller among of new saliva vs old saliva, or that the cfos activation patterns are different in these two conditions.

The answers to these concerns are provided in the Public review Comment from Reviewer #1.

(2) The other key question is if different VMH populations are activated by new vs old saliva.

The answer to this concern is provided in the Public Review comment from Reviewer #1.

Lastly, although the new analysis and text changes improved the manuscript, many issues raised were addressed with some variation of 'future studies will be done', or 'we concur with the Reviewer'. However, the extra experiments required to answer these questions were not done. For this reason, even though the authors have numerous exciting pieces of data, overall the work is still incomplete. I highlight below some examples in which the authors agree with the Reviewer, but do not answer the question with the new work that would be required, or propose to do the work in future studies.

In this revised manuscript, we have conducted several additional experiments to address key concerns raised by the reviewers that are directly relevant to our claims. Specifically, we have examined: 1) whether qualitative or quantitative differences between fresh and old cat saliva are detected by mice to modulate behavior (NEW Figure 1I, J, K, and L, and NEW Supplemental Figure 2); 2) the involvement of Fel d 4 in freezing behavior (NEW Supplemental Figure 5); and 3) whether different VMH populations are activated by fresh versus old saliva (NEW Figure 7). However, some concerns raised by the reviewers fall outside the scope of the current manuscript. These include: 1) identifying the specific components that induce freezing, 2) examining the necessity of V2R-A4 receptors, 3) conducting neural circuit perturbations, and 4) performing a comprehensive analysis—including varying time points, different types of defensive behaviors, and broader neural activation patterns (e.g., cortical versus amygdala activation)—of the mouse’s defensive response to different levels of predator threat imminence. As these aspects are beyond the focus of our current manuscript, we have noted in the Public Review comments.

References:

Allen WE, DeNardo LA, Chen MZ, Liu CD, Loh KM, Fenno LE, Ramakrishnan C, Deisseroth K, Luo L. 2017. Thirst-associated preoptic neurons encode an aversive motivational drive. Science 357:1149– 1155.

DeNardo LA, Liu CD, Allen WE, Adams EL, Friedmann D, Fu L, Guenthner CJ, Tessier-Lavigne M, Luo L. 2019. Temporal evolution of cortical ensembles promoting remote memory retrieval. Nat Neurosci 22:460–469.

Haga-Yamanaka S, Ma L, He J, Qiu Q, Lavis LD, Looger LL, Yu CR. 2014. Integrated action of pheromone signals in promoting courtship behavior in male mice. Elife 3:e03025.

Kimoto H, Haga S, Sato K, Touhara K. 2005. Sex-specific peptides from exocrine glands stimulate mouse vomeronasal sensory neurons. Nature 437:898–901.

Rocha A, Nguyen QAT, Haga-Yamanaka S. 2024. Type 2 vomeronasal receptor-A4 subfamily: Potential predator sensors in mice. Genesis 62:e23597.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Animals have the innate ability to select optimal defensive behavioral outputs with an appropriate intensity in response to predator threat in specific contexts. Such innate behavioral decisions are thought to be computed in the medial hypothalamic nuclei that contain neural populations directly controlling defensive behavioral outputs. The vomeronasal organ (VNO) is one of the major sensory input channels through which predator cues are detected with ascending inputs to the medial hypothalamic nuclei, especially to the ventromedial hypothalamus (VMH), through the medial amygdala (MeA) and bed nucleus of the stria terminalis (BNST). Here, we show that cat saliva contains predator cues that signal imminence of predator threat and regulate the robustness of freezing behavior through the VNO in mice. Cat saliva activates neurons expressing the V2R-A4 subfamily of sensory receptors, suggesting the existence of specific receptor groups responsible for freezing behavior induced by the predator cues. The number of VNO neurons activated in response to saliva correlates with the freshness of saliva and the intensity of freezing behavior, while the downstream neurons in the accessory olfactory bulb (AOB) and defensive behavioral circuit are quantitatively equally activated by fresh and old saliva. Strikingly, however, only the number of VMH neurons activated by fresh saliva positively correlates with the intensity of freezing behavior. Detailed analysis of the spatial distribution of fresh and old saliva-responding neurons revealed a neuronal population within the VMH that is more sensitive to fresh saliva than old saliva. Taken together, this study demonstrates that predator cues in cat saliva change over time and differentially activate the sensory-to-hypothalamus defensive behavioral pathway to modulate behavioral outputs.

Revised version: Animals have the innate ability to select optimal defensive behaviors with appropriate intensity in response to predator threats within specific contexts. Such innate behavioral decisions are thought to be computed in the medial hypothalamic nuclei, which contain neural populations that directly control defensive behavioral outputs. The vomeronasal organ (VNO) serves as a primary sensory channel for detecting predator cues by relaying signals to the medial hypothalamic nuclei, particularly the ventromedial hypothalamus (VMH), via the medial amygdala (MeA) and bed nucleus of the stria terminalis (BNST). Here, we demonstrate that cat saliva contains predator cues that signal the imminence of predator threat and modulate the intensity of freezing behavior through the VNO in mice. Cat saliva activates neurons expressing the V2R-A4 subfamily of sensory receptors, suggesting that specific receptor groups are responsible for inducing the freezing behavior. The number of VNO neurons activated in response to saliva correlates with both the freshness of saliva and the intensity of freezing behavior. In contrast, the downstream neurons in the accessory olfactory bulb (AOB) and the defensive behavioral circuit are activated to a similar extent by fresh and old saliva. Strikingly, however, the number of VMH neurons activated by fresh, but not old, saliva positively correlates with the intensity of freezing behavior. Detailed analysis of the spatial distribution of neurons responding to fresh and old saliva, as well as the overlap of those activated within the same individual mice, revealed that fresh and old saliva predominantly activate distinct neuronal populations within the VMH. Collectively, this study suggests that there is an accessory olfactory circuit in mice that is specifically tuned to time-sensitive components of cat saliva, which optimizes their defensive behavior to maximize their chance of survival according to the imminence of threat.

Introduction

Earlier version: Prey animals make defensive behavioral decisions through prompt and careful analysis of their environment when encountering predators. Immediacy of the threat, determined by perceived distance from predators, is one of the key factors influencing defensive behavioral decisions, which regulates the intensity of behavioral outputs (Figure 1) (Blanchard and Blanchard, 1989). Predator-derived sensory signals, or predator cues, are detected through multiple sensory systems, including the visual system, auditory system, and olfactory system. These individual sensory cues are sufficient to induce innate defensive behavior such as fear-related freezing and flight, as well as anxiety-related risk assessment behavior. In nocturnal animals including rodents, olfaction is one of the major sensory modalities by which predator-derived chemical cues are detected (Apfelbach et al., 2005). Rodents respond to a variety of predator-derived chemical cues, such as odors from fox feces, snake skin, cat fur, and cat collars, in both the field and laboratory environments (Apfelbach et al., 2005; Kats and Dill, 1998; Papes et al., 2010; Pérez-Gómez et al., 2015; Takahashi, 2014). However, how animals use chemosensory cues to make appropriate behavioral decisions is not well understood.

Revised version: Prey animals make defensive behavioral decisions through prompt and careful analysis of their environment when encountering predators. Immediacy of the threat, determined by perceived distance from predators, is one of the key factors influencing defensive behavioral decisions, which regulates the intensity of behavioral outputs (Figure 1A) (Blanchard and Blanchard, 1989). Predator-derived sensory signals, or predator cues, are detected through multiple sensory systems, including the visual system, auditory system, and olfactory system. These individual sensory cues are sufficient to induce innate defensive behavior such as fear-related freezing and flight, as well as anxiety-related risk assessment behavior. In nocturnal animals including rodents, olfaction is one of the major sensory modalities by which predator-derived chemical cues are detected (Apfelbach et al., 2005). Rodents respond to a variety of predator-derived chemical cues, such as odors from fox feces, snake skin, cat fur, and cat collars, in both the field and laboratory environments (Apfelbach et al., 2005; Kats and Dill, 1998; Papes et al., 2010; Pérez-Gómez et al., 2015; Takahashi, 2014). However, how animals use chemosensory cues to make appropriate behavioral decisions is not well understood.

Predator cues in cat saliva induce context-dependent defensive behavior in mice

(A) A schematic diagram illustrating the elicitation of “freeze” or “flight” defensive behaviors in relation to distance. (Adapted with permission from (Blanchard and Blanchard, 1989))

(B) A schematic illustration depicting a test chamber for evaluating defensive behaviors.

(C) Raster plots displaying freezing episodes of individual mice exposed to a clean control swab (gray, n = 18), fresh saliva (green, n = 18), or old saliva (yellow, n = 15). The introduction of a clean, fresh, or old saliva swab into a mouse’s cage is denoted as time 0. Red lines indicate the first contact with a swab for each mouse.

(D) Latency to the first episode of contact with a swab. One-way ANOVA (F (2,47) = 5.409, p = 0.0077) with Tukey’s multiple comparisons test (Ctrl vs. Fresh, p = 0.0640; Ctrl vs Old, p = 0.5932; and Fresh vs Old, p = 0.0074).

(E) The percentage of total freezing episodes after the first contact. One-way ANOVA (F (2,47) = 37.38, p < 0.0001) with Tukey’s multiple comparisons test (Ctrl vs. Fresh, p < 0.0001; Ctrl vs Old, p = 0.0411; and Fresh vs Old, p < 0.0001).

(F) The percentage of total stretch sniff episodes. One-way ANOVA (F(2,49) = 0.007016, p = 0.9930).

(G) The percentage of total direct investigation episodes. One-way ANOVA (F (2,48) = 25.17, p < 0.0001) with Tukey’s multiple comparisons test (Ctrl vs. Fresh, p < 0.0001; Ctrl vs Old, p < 0.0001; and Fresh vs Old, p = 0.5762).

(H) Plasma ACTH concentrations following exposure to clean (n = 10), fresh (n = 8), or old saliva (n = 10) swabs. One-way ANOVA (F (2,25) = 3.735, p = 0.0381) with Tukey’s multiple comparisons test (Ctrl vs. Fresh, p = 0.0370; Ctrl vs Old, p = 0.1609; and Fresh vs Old, p = 0.6804).

Added in the revised version: (I) Volume of saliva in individual swabs at 0 hours (Fresh) and 4 hours (Old) after collection. Two-tailed t-test (t = 1.415, df = 8, p = 0.1949).

Added in the revised version: (J) Concentrations of total proteins in fresh and old saliva. Two-tailed t-test (t = 0.3252, df = 10, p = 0.7518).

Added in the revised version: (K) The percentage of total freezing episodes induced by diluted fresh cat saliva (70%, n = 4; 50%, n = 9; 30%, n = 8), compared to the ones induced by fresh and old saliva. One-way ANOVA (F (4, 47) = 0.6604, p < 0.0001) with Tukey’s multiple comparisons test. Statistical results are reported in Supplemental Table 1.

Added in the revised version: (L) The percentage of total direct investigation episodes to diluted fresh cat saliva (70%, n = 4; 50%, n = 9; 30%, n = 8), compared to the ones induced by fresh and old saliva. One-way ANOVA (F (4, 49) = 2.272, P = 0.0748).

Earlier version: In (D) - (H), the values are presented as means ± S.E.M., with individual dots representing individual mice. *, **, and **** denote significance levels of p < 0.05, 0.01, and 0.0001, respectively.

Revised version: In (D) - (L), the values are presented as means ± S.E.M., with individual dots representing individual mice. *, **, and **** denote significance levels of p < 0.05, 0.01, and 0.0001, respectively. Data for fresh and old saliva in (K) and (L) are reused from (E) and (G), respectively.

The accessory olfactory system is one of the major sensory input channels through which predator-derived chemical cues are detected. It has been shown that some molecules derived from predator species activate sensory neurons in the VNO and induce defensive behavior (Isogai et al., 2011; Papes et al., 2010; Pérez-Gómez et al., 2015; Samuelsen and Meredith, 2009; Tsunoda et al., 2018). Moreover, a number of studies have revealed main components of the sensory neural circuits that mediate behavioral responses to chemical predator cues in rodents (Blanchard et al., 2003; Canteras, 2002; Choi et al., 2005; Dielenberg and McGregor, 2001; Markham et al., 2004; Petrovich et al., 2001). Predator signals detected by the VNO are transferred to the accessory olfactory bulb (AOB). The output neurons in the AOB (the mitral cells) then send the signals to the BNST and the posteroventral part of MeA (MeApv), which connects with the dorsomedial part of the VMH (VMHdm) that contains neural populations directly controlling defensive behavioral outputs (Gross and Canteras, 2012; LeDoux, 2012; Takahashi, 2014). It is, therefore, reasonable to hypothesize that the VNO detects predator cues and sends sensory signals to the VMH, triggering appropriate defensive behavioral decisions in rodents. However, how the sensory signals detected through the VNO-to-VMH circuitry modulate behavioral decisions in specific contexts remains elusive.

In this study, we investigated sensory-to-hypothalamic neural circuitries that regulate defensive behavior induced by non-volatile chemical cues in cat saliva with different freshness. We examined the role of the VNO in behavioral responses and sought sensory receptor genes for the predator cues. Moreover, by comparing downstream neural populations activated by cat saliva with different freshness, we investigated how the imminence of predator threat is processed in the accessory olfactory circuitry to regulate behavioral outputs.

Results

Predator cues in cat saliva induce context-dependent defensive behavior in mice

Earlier version: To explore whether mice exhibit different intensities of defensive behavior in response to chemosensory predator cues with varying threat immediacy, we examined their response to cat saliva with different degrees of freshness, which potentially represents immediacy of the predator threat. Cat saliva has been considered as a source of predator cues found on cat fur and collars, which induce defensive behaviors in rodents (Engelke et al., 2021; Papes et al., 2010). Fresh saliva is assumed to indicate a more immediate threat to the prey animal, which likely triggers stronger defensive behavior. We exposed mice to a cotton swab containing cat saliva in their home cage where they know a flight route is not available (Figure 1B). We found that mice approached and sniffed the cotton swab, regardless of the presence of saliva, as part of their exploratory behavior towards a novel object in their home cage (Supplemental movies 1-3, Figure 1C; red boxes indicate the first contact). The latency to the first contact was slightly longer for saliva collected immediately before introduction to the mouse cage (fresh saliva) when compared to that for saliva collected a few hours prior to the behavior test (old saliva) (Figure 1D). The difference in latencies was statistically significant, indicating that mice can perceive the qualitative change of cat saliva over time (Figure 1D). Baseline freezing episodes were observed in all groups of mice when the sample was introduced into the cage due to the experimenter’s intervention (Figure 1C). However, the groups exposed to cat saliva continued to exhibit freezing behavior throughout the 10-minute observation period, while mice exposed to a clean swab did not exhibit freezing behavior after the first contact (Figure 1C, E). Moreover, we found a correlation between the freshness of cat saliva and duration of freezing behavior in mice. Mice exposed to fresh saliva exhibited considerably longer freezing episodes, accounting for over 50% of the total duration of observation, whereas those exposed to old saliva exhibited freezing episodes for less than 20% of the total duration on average (Figure 1E). We found no significant difference in the latency to the first freezing episode from the first contact with a swab between mice exposed to fresh and old saliva (Supplemental Figure 1A). Mice also exhibited stretch sniffing risk assessment behavior, which did not differ toward any of the swab types (Figure 1F, Supplemental Figure 1B). On the contrary, mice exposed to a swab containing either fresh or old saliva significantly avoided interacting with the swab, while mice exposed to a clean control swab spent a significant amount of time directly interacting with the swab to investigate it, including direct sniffing and chewing (Figure 1G). The comparison of temporal behavioral patterns revealed a slightly higher frequency of direct investigation behavior toward old saliva compared to fresh saliva at the beginning of the exposure period (Supplemental Figure 1C). Importantly, the stress hormone adrenocorticotropic hormone (ACTH) in plasma showed a trend towards two-fold upregulation after exposure to both fresh and old saliva (Figure 1H). These results suggest that mice can still recognize predator threat through old saliva, but freezing behavior in response to fresh saliva attenuates quickly over time.

Revised version: To explore whether mice exhibit different intensities of defensive behavior in response to chemosensory predator cues with varying threat immediacy, we examined their response to cat saliva with different degrees of freshness, which potentially represents immediacy of the predator threat. Cat saliva has been considered as a source of predator cues found on cat fur and collars, which induce defensive behaviors in rodents (Engelke et al., 2021; Papes et al., 2010). Fresh saliva is assumed to indicate a more immediate threat to the prey animal, which likely triggers stronger defensive behavior. We exposed mice to a cotton swab containing cat saliva in their home cage where they know a flight route is not available (Figure 1B). We found that mice approached and sniffed the cotton swab, regardless of the presence of saliva, as part of their exploratory behavior towards a novel object in their home cage (Supplemental movies 1-3, Figure 1C; red boxes indicate the first contact). The latency to the first contact was slightly longer for saliva collected immediately before introduction to the mouse cage (fresh saliva) when compared to that for saliva collected a few hours prior to the behavior test (old saliva) (Figure 1D). The difference in latencies was statistically significant, indicating that mice can perceive the change of cat saliva over time (Figure 1D). Baseline freezing episodes were observed in all groups of mice when the sample was introduced into the cage due to the experimenter’s intervention (Figure 1C). However, the groups exposed to cat saliva continued to exhibit freezing behavior throughout the 10-minute observation period, while mice exposed to a clean swab did not exhibit freezing behavior after the first contact (Figure 1C, E). Moreover, we found a correlation between the freshness of cat saliva and duration of freezing behavior in mice. Mice exposed to fresh saliva exhibited considerably longer freezing episodes, accounting for over 50% of the total duration of observation, whereas those exposed to old saliva exhibited freezing episodes for less than 20% of the total duration on average (Figure 1E). We found no significant difference in the latency to the first freezing episode from the first contact with a swab between mice exposed to fresh and old saliva (Supplemental Figure 1A). Mice also exhibited stretch sniffing risk assessment behavior, which did not differ toward any of the swab types (Figure 1F, Supplemental Figure 1B). On the contrary, mice exposed to a swab containing either fresh or old saliva significantly avoided interacting with the swab, while mice exposed to a clean control swab spent a significant amount of time directly interacting with the swab to investigate it, including direct sniffing and chewing (Figure 1G). The comparison of temporal behavioral patterns revealed a slightly higher frequency of direct investigation behavior toward old saliva compared to fresh saliva at the beginning of the exposure period (Supplemental Figure 1C). Additionally, the stress hormone adrenocorticotropic hormone (ACTH) in plasma showed a trend towards two-fold upregulation after exposure to both fresh and old saliva (Figure 1H). These results suggest that mice can still recognize predator threat through old saliva, but freezing behavior in response to fresh saliva attenuates quickly over time.

Added in the revised version: Saliva-containing swabs were kept in sealed containers immediately after collection from cats to prevent evaporation, ensuring that the saliva volumes at the time of behavioral testing were indistinguishable between fresh and old samples (Figure 1I). The concentrations of total proteins in both fresh and old saliva were also indiscernible, which was approximately 1.5 μg/μL (Figure 1J). Furthermore, the amount of Fel d 4, one of the most abundant proteins in saliva, was nearly equivalent between fresh and old saliva (Supplemental Figure 2), indicating that the differences between fresh and old cat saliva lie in specific components rather than the total or major saliva content. One possible explanation for this difference is the time-dependent reduction of freezing-inducing components in old saliva. Indeed, exposure to diluted fresh saliva resulted in a shorter duration of freezing behavior. Fresh saliva diluted to 70% induced freezing behavior for a duration equivalent to that of undiluted fresh saliva, while freezing behavior in response to 50% and 30% fresh saliva was significantly reduced to the same duration as that observed with old saliva (Figure 1K). The duration of direct interaction with swabs containing 70% and 50–30% fresh saliva also exhibited a similar trend to that observed with fresh and old saliva swabs, respectively (Figure 1L).

Earlier version: Taken together, this behavior study indicates that cat saliva contains predator cues that signal the imminence of predator threat and differentially induce defensive behavioral outputs in mice. More specifically, fresh saliva is recognized as an immediate danger and induces more robust freezing behavior, while old saliva is recognized as a less immediate danger and induces less robust freezing behavioral responses.

Revised version: Taken together, these results indicate that cat saliva contains predator cues that attenuate over time, which is detected by mice to assess the imminence of predator threat and modulate defensive behavioral outputs. More specifically, fresh saliva is recognized as an immediate danger and induces more robust freezing behavior, while old saliva is recognized as a less immediate danger and induces less robust freezing behavioral responses.

Cat saliva activates VNO neurons and modulates freezing behavior

As the freezing behavior was induced when mice made direct contact with the saliva-containing swab, we postulated that direct contact with the predator cues in cat saliva is necessary for freezing behavior induction. To examine the necessity of direct contact, we placed a cotton swab with or without cat saliva into a porous container and investigated mouse’s behavior toward the container. The pores on the container allow volatile chemicals in cat saliva to diffuse out, while nonvolatile chemicals stay inside of the container with the swab. Interestingly, we found that when we introduced a container containing a swab with fresh saliva into the home cage, the mice no longer exhibited freezing behavior towards the container, but instead directly investigated it similarly to the control swab-containing container (Figure 2A, B). This suggests that freezing behavior requires direct contact with the saliva-containing swab and that the freezing-inducing cues in cat saliva detected in this scheme are nonvolatile.

Cat saliva activates VNO neurons and modulates freezing behavior

(A) The percentage of total freezing episodes directed towards control (n = 12) or fresh saliva (n = 11) swabs in cassettes that prevent direct contact with the swabs. Two-tailed t test (t = 0.08745, df = 21, p = 0.9311).

(B) The percentage of total investigation episodes toward cassettes containing control or fresh saliva swabs. Two-tailed t test (t=1.662, df=21, p = 0.1114).

(C) The percentage of total freezing episodes directed toward control or fresh saliva swabs in Trpc2-WT (n = 10) or -KO (n = 9) mice. Two-way ANOVA ([genotype] F (1, 34) = 3.662, p = 0.0641, [swab type] F (1, 34) = 8.053, p = 0.0076, [genotype] and [swab type] F (1, 34) = 5.963, p = 0.0200) with Tukey’s multiple comparisons test (WT-Ctrl vs KO-Ctrl, p = 0.9922; WT-Ctrl vs WT-Saliva, p = 0.0164; WT-Ctrl vs KO-Saliva, p = 0.9136; KO-Ctrl vs WT-Saliva, p = 0.0100; KO-Ctrl vs KO-Saliva, p = 0.9832; WT-Saliva vs KO-Saliva, p = 0.0037).

(D) The percentage of total direct investigation episodes toward control or fresh saliva swabs in Trpc2-WT or -KO mice. Two-way ANOVA ([genotype] F (1, 34) = 11.46, p = 0.0018, [swab type] F (1, 34) = 7.725, p = 0.0088, [genotype] and [swab type] F (1, 34) = 13.96, p = 0.0007) with Tukey’s multiple comparisons test (WT-Ctrl vs KO-Ctrl, p = 0.9051; WT-Ctrl vs WT-Saliva, p < 0.0001; WT-Ctrl vs KO-Saliva, p = 0.9732; KO-Ctrl vs WT-Saliva, p = 0.0006; KO-Ctrl vs KO-Saliva, p = 0.9949; WT-Saliva vs KO-Saliva, p = 0.0003).

(E) Representative images illustrating the expression of pS6 (green) and Gαo (magenta) in the VNO of mice exposed to control, fresh, or old saliva swabs. DAPI was used as a nuclear counterstain. Scale bar: 50 μm.

(F, G) The number of pS6-IR positive cells in 1000 VNO cells in the apical (F) and basal (G) regions of the VNO neuroepithelium in mice stimulated with control (n = 8), fresh saliva (n = 8), and old saliva (n = 8) swabs. (F) One-way ANOVA (F (2, 21) = 2.562, p = 0.1010). (G) One-way ANOVA (F (2, 21) = 17.90, p < 0.0001) with Tukey’s multiple comparisons test (Ctrl vs. Fresh, p < 0.0001; Ctrl vs Old, p = 0.2262; and Fresh vs Old, p = 0.0014).

(H) The percentage of total freezing episodes in relation to the number of pS6-positive cells in the VNO of individual mice (n = 8 per swab type). Correlation analysis was conducted using Spearman’s rank correlation coefficient as total data points did not pass Shapiro-Wilk and Kolmogorov-Smirnov normality tests. Statistical results are reported in the graph.

(A - D, F, G) the values are presented as means ± S.E.M., with individual dots representing individual mice. *, **, ***, and **** denote significance levels of p < 0.05, 0.01, 0.001, and 0.0001, respectively.

Mice have the ability to detect nonvolatile chemical cues through the VNO (Mohrhardt et al., 2018). Ethologically relevant chemical cues from different animals, such as pheromones and predator cues, are sucked up into the VNO and activate VNO sensory neurons when animals contact the chemical source (Luo et al., 2003). We, therefore, explored the role of the VNO in inducing freezing in response to cat saliva using a VNO-deficient laboratory mouse line that lacks the primary signal transduction ion channel in VNO neurons (transient receptor potential cation channel, subfamily C, member 2-knockout line or Trpc2-KO line) (Leypold et al., 2002). In Trpc2-KO mice, chemosensory signals detected by sensory receptor proteins are expected to be attenuated before being relayed to the brain, due to the lack of the downstream transduction channel in VNO neurons (Leypold et al., 2002; Liman et al., 1999; Stowers et al., 2002). Thus, Trpc2-KO mice have been widely utilized to examine roles of the VNO in various behaviors (Beny-Shefer et al., 2017; Ferrero et al., 2013; Haga et al., 2010; Kimchi et al., 2007; Leypold et al., 2002; Papes et al., 2010; Pérez-Gómez et al., 2015; Stowers et al., 2002). Upon exposure to a fresh saliva-containing swab, Trpc2-KO mice failed to exhibit freezing behavior, in contrast to their wild-type littermates who displayed freezing responses. Additionally, the Trpc2-KO mice directly investigated the saliva swab with a similar level of intensity as the control swab, while the wild-type littermates actively avoided it (Figure 2C, D). These findings suggest that VNO activation, disrupted in Trpc2-KO mice, is necessary for the freezing behavior observed in response to cat saliva in mice.

To investigate if VNO neurons are indeed activated by cat saliva, we analyzed immunoreactivity (IR) of phosphorylated ribosomal protein S6 subunit, pS6, in the VNO of mice exposed to cat saliva. As the cellular pS6 level is greatly upregulated upon activation of VNO neurons, it has been used as a sensitive marker to examine neural activation in the VNO (Carvalho et al., 2020; Itakura et al., 2022; Osakada et al., 2022; Tsunoda et al., 2018). We observed that exposure to both fresh and old cat saliva induced an increase in pS6-IR in the basal VNO neurons that express Gαo protein (Figure 2E, G). Conversely, pS6-IR was rarely observed in the apical VNO neurons (Figure 2F). Notably, we found that the number of pS6-positive neurons was more than three times higher in mice exposed to fresh saliva compared to that exposed to old saliva (Figure 2G). Strikingly, when the duration of freezing was plotted as a function of pS6-IR signals in the VNO of individual animals, we observed a significantly positive correlation (R2 = 0.3586, 95% CI [0.2458, 0.8116], p = 0.0020) between the numbers of active sensory neurons and behavioral outputs (Figure 2H), suggesting that more intense freezing behavior is induced when more VNO neurons are activated by cat saliva.

Taken together, these results indicate that freezing-inducing predator cues are detected by Gαo-positive sensory neurons in the VNO, and these sensory neurons convey the imminence of predator threat to higher brain centers.

Cat saliva activates the V2R-A4 subfamily of the vomeronasal receptors

The significant reduction of the freezing duration to aged saliva reflects possible attenuation of signal inputs to a specific population of sensory receptors, or alternatively, it may reflect a change of sensory receptor repertoires due to qualitative changes of predator cues over time. To test these possibilities, we examined sensory receptors expressed in the neurons activated by fresh and old saliva. The basal VNO neurons express sensory receptors that belong to vomeronasal receptor type 2 (V2R) family of G protein-coupled receptors in mice (Figure 3A) (Mohrhardt et al., 2018). The mouse V2R family consists of over 120 intact genes, which are divided into V2R A1-6, A8-10, and B-D subfamilies based on sequence homology (Figure 3A) (Francia et al., 2014; Silvotti et al., 2011, 2007). With the exception of the V2R C group that is co-expressed with other groups of V2Rs in the basal VNO neurons, each sensory neuron expresses only one member of the V2R receptor, allowing individual neurons to respond to specific molecules (Mohrhardt et al., 2018).

Cat saliva activates the V2R-A4 subfamily of the vomeronasal receptors

(A) Phylogenetic tree of the V2R family and its subfamilies used for V2R cRNA probe design.

(B) Representative images displaying the expression of V2RA4 (magenta) and pS6 (green) in the VNO of a mouse exposed to a fresh saliva swab. DAPI was utilized as a nuclear counterstain. Scale bar: 50 μm.

(C) The percentage of cells co-labeled with V2R probes and anti-pS6 antibody among the total pS6-positive cells in the VNO of mice exposed to fresh saliva (top, from 3 – 6 animals) or old saliva (bottom, from 3 – 6 animals).

Earlier version: To investigate which subfamilies of V2R are expressed in cat saliva-activated VNO neurons, we designed in situ hybridization (ISH) probes that detect individual V2R subfamilies and performed dual staining of V2R ISH with pS6-IR (Supplemental Figure 2). We found that the majority of pS6-IR signals induced by cat saliva co-localized with signals of the ISH probe against the V2R-A4 subfamily (Figure 3B, Supplemental Figure 2). Approximately 45% of pS6-IR signals induced by fresh saliva (83 out of 181) co-localized with V2R-A4 probe signals, with less than 1 % of pS6-IR (1 out of 120) co-localized with A8-4 probe signals (Figure 3C). Notably, V2R-A4 probe signals also co-localized with approximately 45% of pS6-IR signals induced by old saliva (37 out of 86). None of the other probe signals co-localized with pS6-IR induced by both fresh and old saliva. These results indicate that members of the V2R-A4 subfamily are potential sensory receptors for the predator cues contained in both fresh and old cat saliva, and varying levels of imminence of predator threat stimulate the same type of sensory receptors in a quantitatively different manner to regulate the freezing behavior.

Revised version: To investigate which subfamilies of V2R are expressed in cat saliva-activated VNO neurons, we designed in situ hybridization (ISH) probes that detect individual V2R subfamilies and performed dual staining of V2R ISH with pS6-IR (Supplemental Figure 3). We found that the majority of pS6-IR signals induced by cat saliva co-localized with signals of the ISH probe against the V2R-A4 subfamily (Figure 3B, Supplemental Figure 3). Approximately 45% of pS6-IR signals induced by fresh saliva (83 out of 181) co-localized with V2R-A4 probe signals, with less than 1 % of pS6-IR (1 out of 120) co-localized with A8-4 probe signals (Figure 3C). Notably, V2R-A4 probe signals also co-localized with approximately 45% of pS6-IR signals induced by old saliva (37 out of 86). None of the other probe signals co-localized with pS6-IR induced by both fresh and old saliva. These results indicate that members of the V2R-A4 subfamily are potential sensory receptors for the predator cues contained in both fresh and old cat saliva, and varying levels of imminence of predator threat stimulate the same type of sensory receptors in a quantitatively different manner to regulate the freezing behavior.

Cat saliva activates neurons in the AOB and defensive behavioral circuit

We next aimed to understand how the differential activation of sensory receptor neurons in the VNO by different freshness of cat saliva influences behavioral decisions through downstream neural circuits (Figure 4A). We investigated this by examining the induction of cFos, an immediate early gene product, in the projection neurons (mitral cells) and local inhibitory neurons (granule cells) in the AOB of mice exposed to cat saliva. The axons of apical and basal VNO neurons project to the anterior and posterior regions of the AOB (aAOB and pAOB), respectively (Mohrhardt et al., 2018). We found that the activation of the basal VNO neurons by cat saliva leads to predominant induction of cFos-IR signals in the pAOB, whose glomerulus layer is Gαi-IR negative (Figure 4B-F). The numbers of cFos-positive mitral and granule cells in the pAOB in saliva-exposed mice are significantly higher than those in mice exposed to the control swab (Figure 4D, F). However, contrary to our expectations, we did not observe any differences in the numbers of cFos-positive mitral and granule cells in the AOB between fresh and old saliva-exposed mice (Figure 4D, F). These results suggest that, although the AOB processes saliva signals from the VNO following predator cue detection, the relationship between the degree of AOB activation and freezing behavior intensity is non-linear.

Cat saliva activates neurons in the AOB and defensive behavioral circuit

(A) Schematic diagrams illustrating the VNO sensory pathway that controls defensive behaviors.

(B) Representative images displaying the expression of cFos (magenta) in the mitral cell (MCL) and granule cell (GCL) layers of the AOB of mice exposed to control, fresh, or old saliva swabs. Gαi2 (green) visualizes the anterior region of the glomerular layer (GL). DAPI was utilized as a nuclear counterstain. Scale bar: 200 μm.

(C, D) The number of cFos-IR positive mitral cells per 1000 mitral cells in the anterior (C) and posterior (D) regions of the AOB (n = 10 per swab type). (C) One-way ANOVA (F (2, 27) = 0.6106, p = 0.5504). (D) One-way ANOVA (F (2, 27) = 6.665, p = 0.0044) with Tukey’s multiple comparisons test (Ctrl vs. Fresh, p = 0.0210; Ctrl vs Old, p = 0.0059; and Fresh vs Old, p = 0.8593).

(E, F) The number of cFos-IR positive granule cells per 1000 granule cells in the anterior (E) and posterior (F) regions of the AOB (n = 10 per swab type). (E) One-way ANOVA (F (2, 27) = 0.5831, p = 0.5651). (F) One-way ANOVA F (2, 27) = 5.996, p = 0.0070) with Tukey’s multiple comparisons test (Ctrl vs. Fresh, p = 0.0262; Ctrl vs Old, p = 0.0099; and Fresh vs Old, p = 0.9105).

(G) The number of cFos-IR positive neurons per 1000 cells in the BNST in mice stimulated with control (n = 10), fresh saliva (n = 5), and old saliva (n = 6) swabs. One-way ANOVA (F (2, 18) = 0.4375, p = 0.6524).

(H) The number of cFos-IR positive neurons per 1000 cells in the MeApv in mice stimulated with control (n = 9), fresh saliva (n = 6), and old saliva (n = 7) swabs. One-way ANOVA (F (2, 19) = 2.454, p = 0.1128).

(I) The number of cFos-IR positive neurons per 1000 cells in the VMH in mice stimulated with control (n = 6), fresh saliva (n = 6), and old saliva (n = 6) swabs. One-way ANOVA (F (2, 15) = 6.368, p = 0.0100) with Tukey’s multiple comparisons test (Ctrl vs. Fresh, p = 0.0116; Ctrl vs Old, p = 0.0374; and Fresh vs Old, p = 0.8247).

(J) The number of cFos-IR positive neurons per 1000 cells in the dPAG in mice stimulated with control (n =5), fresh saliva (n = 6), and old saliva (n = 8) swabs. One-way ANOVA (F (2, 16) = 2.136, p = 0.1506).

In (C-J), the values are presented as means ± S.E.M., with individual dots representing individual mice. * and ** denote significance levels of p < 0.05 and 0.01, respectively.

Earlier version: To further investigate how signals of fresh and old saliva are processed in the downstream circuitry, we examined induction of cFos in the neural substrates involved in innate defensive behaviors (Figure 4A). We investigated the BNST and MeApv, as these nuclei receive direct ascending inputs from the mitral cells in the AOB (Mohrhardt et al., 2018). We also investigated the VMH and dPAG, as these nuclei have been shown to control defensive behavioral outputs (Silva et al., 2013). Upon exposure to fresh and old saliva, cFos expression in the MeApv, VMH, and dPAG, but not in the BNST, appeared to be upregulated compared to the control stimulus (Figure 4G-J). Notably, cFos expression was significantly increased in the VMH of mice exposed to fresh and old saliva compared to the control (Figure 4I), consistent with previous observations that the dorsomedial and central subdivisions of VMH (VMHdm/c) contain subpopulations of neurons that control freezing behavior (Engelke et al., 2021; Kunwar et al., 2015; Silva et al., 2016, 2013; Wang et al., 2015) However, there was no difference in the number of cFos-positive cells between the two saliva-exposed groups in the neural substrates tested (Figure 4G-J). These results suggest that, although the defensive circuit is activated by cat saliva signals, the relationship between the degree of activation and freezing behavior intensity is non-linear, similar to the situation in the AOB.

Revised version: To further investigate how signals of fresh and old saliva are processed in the downstream circuitry, we examined induction of cFos in the neural substrates involved in innate defensive behaviors (Figure 4A). We investigated the BNST and MeApv, as these nuclei receive direct ascending inputs from the mitral cells in the AOB (Mohrhardt et al., 2018). We also investigated the VMH and dPAG, as these nuclei have been shown to control defensive behavioral outputs (Silva et al., 2013). Upon exposure to fresh and old saliva, cFos expression in the MeApv, VMH, and dPAG, but not in the BNST, appeared to be upregulated compared to the control stimulus (Figure 4G-J, Supplemental Figure 4). Notably, cFos expression was significantly increased in the VMH of mice exposed to fresh and old saliva compared to the control (Figure 4I), consistent with previous observations that the dorsomedial and central subdivisions of VMH (VMHdm/c) contain subpopulations of neurons that control freezing behavior (Engelke et al., 2021; Kunwar et al., 2015; Silva et al., 2016, 2013; Wang et al., 2015). However, there was no difference in the number of cFos-positive cells between the two saliva-exposed groups in the neural substrates tested (Figure 4G-J). These results suggest that, although the defensive circuit is activated by cat saliva signals, the relationship between the degree of activation and freezing behavior intensity is non-linear, similar to the situation in the AOB.

Taken together, these results indicate that the relationship between neural activity and behavioral output is more complex in the downstream circuit than in the VNO. More specifically, the numbers of active neurons and intensity of freezing behavior do not appear to correlate in the downstream regions. This suggests that fresh and old cat saliva may differentially activate functionally distinct neural populations within these regions, which in turn leads to quantitatively different freezing behavioral responses.

Fresh and old cat saliva differentially activate VMH neurons to control defensive behavior

To investigate the potential differences between fresh and old saliva-activated neurons, we analyzed correlations of the duration of freezing behavior and numbers of cFos-IR signals in different cell types in individual animals from different exposure groups. To our surprise, we found that the numbers of cFos-IR signals in AOB mitral and granule cells, as well as those in neurons in the BNST and dPAG, did not correlate with the duration of freezing behavior in any of the exposure groups (Figure 5A-C, F). However, we observed a significant positive correlation between the number of cFos-IR signals and behavior intensity only in the fresh saliva-exposed group in the VMH (R2 = 0.5708, 95% CI [-0.1449, 0.9714], p = 0.0412) while no such correlation was observed in the old saliva-exposed group (Figure 5E). We observed a similar positive correlation trend in the MeApv (R2 = 0.3854, 95% CI [-0.3845, 0.9525], p = 0.0942), although it was not statistically significant possibly due to low sample numbers (Figure 5D). Taken together, these findings support our hypothesis that fresh and old saliva may differently activate the downstream circuit.

Fresh and old cat saliva differentially activate VMH neurons to control defensive behavior

(A-F) Percentage of the total freezing episodes as a function of the number of cFos-positive posterior AOB mitral cells (A) and granule cells (B), BNST (C), MeApv (D), VMH (E), and dPAG (F) neurons. Green, and yellow symbols represent values from mice exposed to fresh saliva and old saliva, respectively. Sample size is same as seen in figure 4. Correlation was analyzed using Pearson correlation coefficient as data points in each group pass Shapiro-Wilk and Kolmogorov-Smirnov normality tests. Statistical results are reported in the table below each graph. * denotes statistical significance with p < 0.05.

Earlier version: We therefore examined whether fresh and old cat saliva activate neurons in spatially distinct regions within the VMH. We observed the distribution of cFos-IR neurons in five serial brain sections from rostral to caudal VMHdm regions. The VMH region within each hemisphere section was further divided into 10 bins along the dorsomedial-ventrolateral axis, and cFos-IR signals in each bin were counted (Figure 6A). The majority of cFos-IR signals induced by both fresh and old cat saliva were localized in the dorsomedial to central part of the VMH (bins 1 – 5). Strikingly, the numbers of cFos-IR signals induced by fresh and old saliva were significantly different in the section 2 bin 5, as well as the section 3 bin 3 and bin 4, which are adjacent to each other (Figure 6B). The cFos-IR signals in the other regions did not show significant difference between fresh and old saliva-exposed mice. These results suggest that there is a distinct population of neurons clustered in a close proximity within the VMH, which is more sensitive to fresh saliva than old saliva. This population of neurons in the VMHdm may play a role in inducing robust freezing behavior to fresh cat saliva. Moreover, these results also suggest that the neurons that induce robust freezing to fresh saliva and less robust freezing to old saliva reside in the same region of the VMH and may use different mechanisms to encode freezing intensity. Overall, our results suggest that the qualitative properties of predator signal inputs, such as their freshness, can induce distinct neural processing and ultimately lead to different defensive behavioral outcomes.

Revised version: We therefore examined whether fresh and old cat saliva activate neurons in spatially distinct regions within the VMH. We observed the distribution of cFos-IR neurons in five serial brain sections from rostral to caudal VMHdm regions. The VMH region within each hemisphere section was further divided into 10 bins along the dorsomedial-ventrolateral axis, and cFos-IR signals in each bin were counted (Figure 6A). The majority of cFos-IR signals induced by both fresh and old cat saliva were localized in the dorsomedial to central part of the VMH (bins 1 – 5). Strikingly, the numbers of cFos-IR signals induced by fresh and old saliva were significantly different in the section 2 bin 5, as well as the section 3 bin 3 and bin 4, which are adjacent to each other (Figure 6B). The cFos-IR signals in the other regions did not show significant difference between fresh and old saliva-exposed mice. These results suggest that there is a distinct population of neurons clustered in a close proximity within the VMH, which is more sensitive to fresh saliva than old saliva. This population of neurons in the VMHdm may play a role in inducing robust freezing behavior to fresh cat saliva. Moreover, these results also suggest that the neurons that induce robust freezing to fresh saliva and less robust freezing to old saliva reside in the same region of the VMH and may use different mechanisms to encode freezing intensity.

A population of VMH neurons is more sensitive to fresh saliva than old saliva

(A) Schematic diagram illustrating the cFos-IR distribution analysis. Each VMH region in five consecutive sections (140 μm apart) was prepared along the dorsomedial (dm)-ventrolateral (vl) axis and divided into 10 bins that are vertical to the dm-vl axis. cFos-IR signals in each bin were counted.

(B) The number of cFos-IR-positive cells in each bin in the VMH hemisphere sections. Sample images are the same as seen in figure 4 (n = 6 – 7 animals). The values are presented as means + S.E.M., with individual dots representing values from individual VMH hemisphere sections. Statistical analysis was performed using Two-Way ANOVA with Tukey’s multiple comparisons. Statistical results are reported in Supplemental Table 1. Black * and ** denote significance levels of p < 0.05 and 0.01, respectively, between control and fresh or old saliva conditions. Red * and ** denote significance levels of p < 0.05 and 0.01, respectively, between fresh and old saliva conditions.

Added in the revised version: Fresh and old cat saliva activate largely different populations of VMH neurons

Added in the revised version: Do fresh and old saliva activate distinct VMH neurons within the same individuals? To investigate this, we performed a double exposure experiment using Fos2A-iCreERT2; Ai9 mice (Figure 7A). The Fos2A-iCreERT2 allele enabled the expression of iCre-ERT2 recombinase in activated neurons, facilitating loxP-mediated recombination in a 4-hydroxytamoxifen (4-OHT)-dependent manner (Allen et al., 2017; DeNardo et al., 2019). This recombination resulted in the expression of tdTomato from the Ai9 allele, allowing visualization of neurons activated during the first saliva exposure. One week after the initial exposure, the same mice were subjected to a second saliva exposure for one hour, leading to cFos expression induced by the neural activity triggered during the second exposure. The cFos protein was detected by immunohistochemistry (Figure 7A).

Added in the revised version: Fresh and old cat saliva activate largely different populations of VMH neurons.

Added in the revised version: (A) Schematic diagram illustrating the double exposure experiment using Fos2A-iCreERT2; Ai9 mice to visualize the populations of neurons activated by fresh and old saliva in the same individual mouse.

Added in the revised version: (B) Representative images showing tdTomato expression (magenta) and cFos-IR (green) in the VMH of mice exposed to different swabs: Control (1st and 2nd), Fresh saliva (1st and 2nd), or Old saliva (1st) and Fresh saliva (2nd). Overlapping signals are indicated in white. DAPI was used for nuclear counterstaining. Yellow dotted lines outline the VMH. Areas enclosed by yellow squares in the left images are enlarged in the right images. Scale bar: 100 μm.

Added in the revised version: (C) Schematic diagram illustrating the color code for the graphs in (D) and (E).

Added in the revised version: (D) The percentage of double-positive cells among tdTomato-expressing cells in mice exposed to Fresh (1st) and Fresh (2nd) saliva (n = 5) and Old (1st) and Fresh (2nd) saliva (n = 5). A permutation test (p = 0.0035).

Added in the revised version: (E) The percentage of double-positive cells among cFos-IR-positive cells in mice exposed to Fresh (1st) and Fresh (2nd) saliva (n = 5) and Old (1st) and Fresh (2nd) saliva (n = 5). A permutation test (p = 0.0060).

Added in the revised version: In (D) and (E), the values are presented as means ± S.E.M., with individual dots representing individual mice. ** denotes a significance level of p < 0.01.

Added in the revised version: In animals exposed to both control (1st) and control (2nd) swabs, sparse tdTomato and cFos-IR signals were observed in the VMH. However, in animals exposed to two saliva samples, substantial populations of neurons in the VMH expressed tdTomato and/or cFos-IR signals (Figure 7B). Quantification of tdTomato and cFos-IR double-positive cells among tdTomato-labeled cells (Figure 7C) revealed that 43% (mean per animal: 61 / 143) of the cells activated by fresh saliva during the first exposure were also activated by fresh saliva during the second exposure, whereas only 16% (17 / 106) of cells activated by old saliva during the first exposure were activated by fresh saliva during the second exposure (p = 7.5e-6, Chi-squared test). The difference in the fraction of overlapping cells between fresh and old saliva exposures was maintained when we compared the two groups of animals (Figure 7D, p = 0.0035, permutation test). Additionally, quantification of tdTomato and cFos-IR double-positive cells among cFos-IR cells indicated that 27% (61 / 226) of cells activated by fresh saliva during the second exposure were previously activated by fresh saliva, whereas 15% (17 / 112) of cells activated by fresh saliva during the second exposure were previously activated by old saliva (p = 0.015, Chi-squared test). The difference in the fraction of overlapping cells between fresh and old saliva exposures was also significant in this analysis (Figure 7E, p = 0.0060, permutation test). Together, these results demonstrate that fresh and old cat saliva activate largely different populations of neurons within the VMH.

Added in the revised version: Overall, our results suggest that the qualitative properties of predator signals, such as freshness, can induce distinct neural processing and ultimately lead to different defensive behavioral outcomes.

Discussion

Earlier version: This study provides evidence for the role of cat saliva as a source of predator cues that signal imminent predator threat and regulate the intensity of freezing behavior through the accessory olfactory pathway in mice (Figure 7). Notably, both fresh and old saliva activate neurons expressing the V2R-A4 subfamily of sensory receptors, suggesting the existence of a specific receptor group responding to the predator cues. Importantly, the number of vomeronasal neurons activated in response to saliva correlates with the freshness of saliva and the intensity of freezing behavior. Moreover, the number of VMH neurons activated by immediate threat (fresh saliva) exhibits a positive correlation with the intensity of behavioral outputs, suggesting their direct involvement in the modulation of freezing behavior. In contrast, the activation of VMH neurons by less immediate threat (old saliva) does not show such a correlation. These contrasting responses to different freshness levels of cat saliva suggest the presence of distinct populations of neurons in the VMH controlling the intensity of freezing behavior. These findings establish a framework for understanding the neural mechanisms underlying innate defensive behavior decisions.

Revised version: This study provides evidence for the role of cat saliva as a source of predator cues that signal imminent predator threat and regulate the intensity of freezing behavior through the accessory olfactory pathway in mice (Figure 8). Notably, both fresh and old saliva activate neurons expressing the V2R-A4 subfamily of sensory receptors, suggesting the existence of a specific receptor group responding to the predator cues. Importantly, the number of vomeronasal neurons activated in response to saliva correlates with the freshness of saliva and the intensity of freezing behavior. Moreover, the number of VMH neurons activated by immediate threat (fresh saliva) exhibits a positive correlation with the intensity of behavioral outputs, suggesting their direct involvement in the modulation of freezing behavior. In contrast, the activation of VMH neurons by less immediate threat (old saliva) does not show such a correlation, suggesting the presence of distinct populations of neurons in the VMH processing imminence of predator threat. Consistent with this, we detected a specific population of VMH neurons that are more sensitive to fresh saliva than to old saliva. We also showed that fresh and old saliva activate largely different populations of VMH neurons in the same individuals, further suggesting the existence of distinct populations of VMH neurons that modulate defensive behavior according to the imminence of threat.

A model of the VMH-to-VMH pathway processing imminence of predator threat

Earlier version: The circuitry is selectively activated based on the varying imminence of predator cues. Fresh saliva activates more V2R-A4-expressing VNO neurons (blue), which results in stronger activation of the mitral cells innervating V2R-A4 glomeruli. This eventually triggers freezing-correlated neural activity in the VMH. In contrast, old saliva activates fewer V2R-A4-expressing VNO neurons (blue), which results in weaker activation of the mitral cells innervating V2R-A4 glomeruli, leading to freezing-non-correlated activity in the VMH.

Revised version: The circuitry is selectively activated based on the varying imminence of predator cues. Fresh saliva activates more V2R-A4-expressing VNO neurons (blue), which potentially results in stronger activation of the V2R-A4 glomeruli. This eventually triggers freezing-correlated neural activity in the VMH. In contrast, old saliva activates fewer V2R-A4-expressing VNO neurons (blue), which potentially results in weaker activation of the V2R-A4 glomeruli, leading to freezing-non-correlated activity in the VMH.

The VNO as the sensor of predator cues that induce fear-related behavior

This study demonstrates that the VNO is involved in detecting and processing chemosensory predator signals that induce fear-related freezing behavior. Previous studies have shown that the VNO mediates defensive behaviors against predators that are rather related to anxiety state, such as risk assessment, area avoidance, and reduced exploratory and locomotion behaviors (Isogai et al., 2011; Papes et al., 2010; Pérez-Gómez et al., 2015; Tsunoda et al., 2018). The difference between these studies and our current study is likely due to different sample types and/or their freshness, as we observed that fear-related behavior is correlated with freshness of predator samples. It is interesting to assume that the biological sources of predator cues, such as saliva, urine, and tears, contain varieties of unidentified cues that convey time-sensitive signals through the mouse accessory olfactory system. This is consistent with the fact that VNO neurons expressing more than half of vomeronasal receptors and a significant fraction of AOB mitral cells have been shown to be activated by predator-derived samples (Ben-Shaul et al., 2010; Isogai et al., 2011). Therefore, it is important to investigate how the sensory pathway in the accessory olfactory system processes varieties of predator cues, which eventually leads to defensive behavioral decisions in higher brain centers.

Earlier version: Moreover, our study demonstrates that time-dependent alterations of the chemosensory predator signals can be detected by the VNO, which induces altered defensive behavioral responses. For instance, 4-hour-old saliva induces reduced freezing behavior compared to fresh saliva. On the other hand, duration of direct investigation remains low toward both fresh and old saliva, and the stress hormone ACTH level is upregulated in both cases. A future study ought to identify specific molecules—most likely proteins or peptides—in cat saliva responsible for inducing these distinct defensive responses in mice. A cat lipocalin allergen, Fel d 4, has been identified to induce risk assessment defensive behavior in mice (Papes et al., 2010). While Fel d 4 is one of the potential candidates, there exists a possibility of a different molecule or a combination of multiple molecules playing a role. Once the molecules are identified, it is imperative to investigate how their quantity and quality change over time and how these factors correlate with freezing behavior. Such an exploration will provide insights into how mice sense the imminence of threat using those chemical signals.

Revised version: Moreover, our study indicates that time-dependent alterations of the chemosensory predator signals can be detected by the VNO, which induces altered defensive behavioral responses. For instance, 4-hour-old saliva induces reduced freezing behavior compared to fresh saliva. On the other hand, duration of direct investigation remains low toward both fresh and old saliva, and the stress hormone ACTH level is upregulated in both cases. A future study ought to identify specific molecules—most likely proteins or peptides—in cat saliva responsible for inducing freezing behavior in mice. A cat lipocalin allergen, Fel d 4, has been identified as a trigger for risk assessment defensive behavior in mice (Papes et al., 2010). While Fel d 4 appears to be a potential candidate for inducing freezing behavior, our analysis showed that the amount of Fel d 4 did not differ between fresh and old saliva (Supplemental Figure 2). Additionally, we found that recombinant Fel d 4 protein did not induce freezing behavior (Supplemental Figure 5). Therefore, it is unlikely that Fel d 4 alone is the factor responsible for triggering freezing behavior. This suggests the possibility of an unidentified molecule or a combination of multiple unidentified molecules contributing to this response. Once these molecules are identified, it will be crucial to investigate how their quantity and quality change over time and how these changes correlate with freezing behavior. Such an investigation will offer valuable insights into how mice perceive the immediacy of a threat through these chemical signals.

V2R-A4 subfamily as the receptor for predator cues in cat saliva

Earlier version: Screening of sensory receptors in the VNO neurons activated by fresh and old saliva demonstrated that members of the V2R-A4 subfamily are the candidate receptors for the predator cues in cat saliva. The V2R-A4 subfamily has been shown to be expressed in the VNO neurons that are activated by biological samples derived from various mammalian predators (Carvalho et al., 2020; Isogai et al., 2011; Tsunoda et al., 2018). Therefore, it is likely that specific members of the V2R-A4 subfamily are the receptors for the predator cues in cat saliva that induce freezing behavior.

Revised version: Screening of sensory receptors in the VNO neurons activated by fresh and old saliva demonstrated that members of the V2R-A4 subfamily are the candidate receptors for the predator cues in cat saliva. The V2R-A4 subfamily has been shown to be expressed in the VNO neurons that are activated by biological samples derived from various mammalian predators (Carvalho et al., 2020; Isogai et al., 2011; Rocha et al., 2024; Tsunoda et al., 2018). Therefore, it is likely that specific members of the V2R-A4 subfamily are the receptors for the predator cues in cat saliva that induce freezing behavior.

Earlier version: Dual staining of vomeronasal receptors with immediate early gene products has been effectively utilized to identify receptors for several pheromones (Carvalho et al., 2020; Haga et al., 2010; Isogai et al., 2018, 2011; Itakura et al., 2022; Kimoto et al., 2005; Osakada et al., 2022, 2018; Tsunoda et al., 2018). However, due to very high sequence homologies among members of the V2R-A4 subfamily (Francia et al., 2014), alternative approaches need to be considered to further isolate the receptors responsible for the detection of predator cues in cat saliva. Potential approaches include receptor mRNA profiling with pS6 immunoprecipitation from the VNO of mice stimulated by cat saliva (Isogai et al., 2018; Jiang et al., 2015), as well as receptor mRNA profiling from isolated single cells activated by cat saliva in GcaMP imaging using the VNO slices in vitro (Haga-Yamanaka et al., 2014; Wong et al., 2020). Receptor candidates identified using either of the methods can be further confirmed by examining necessity and sufficiency for detecting cat saliva using genetically modified mouse lines (Haga et al., 2010; Haga-Yamanaka et al., 2014; Isogai et al., 2018; Itakura et al., 2022; Osakada et al., 2022, 2018; Tsunoda et al., 2018).

Revised version: Dual staining of vomeronasal receptors with immediate early gene products has been effectively utilized to identify receptors for several pheromones (Carvalho et al., 2020; Haga et al., 2010; Isogai et al., 2018, 2011; Itakura et al., 2022; Kimoto et al., 2005; Osakada et al., 2022, 2018; Tsunoda et al., 2018). However, due to very high sequence homologies among members of the V2R-A4 subfamily (Francia et al., 2014; Rocha et al., 2024), alternative approaches need to be considered to further isolate the receptors responsible for the detection of predator cues in cat saliva. Potential approaches include receptor mRNA profiling with pS6 immunoprecipitation from the VNO of mice stimulated by cat saliva (Isogai et al., 2018; Jiang et al., 2015), as well as receptor mRNA profiling from isolated single cells activated by cat saliva in GCaMP imaging using the VNO slices in vitro (Haga-Yamanaka et al., 2014; Wong et al., 2020). Receptor candidates identified using either of the methods can be further confirmed by examining necessity and sufficiency for detecting cat saliva using genetically modified mouse lines (Haga et al., 2010; Haga-Yamanaka et al., 2014; Isogai et al., 2018; Itakura et al., 2022; Osakada et al., 2022, 2018; Tsunoda et al., 2018).

Differential processing of fresh and old saliva signals in the VNO-to-VMH pathway

Our correlation analyses between the number of activated neurons and behavior intensity suggest that sensory signals encoding different imminence of predator threat are differentially processed in the VNO-to-VMH pathway. The differential behavioral responses do not appear to result from global neuromodulations of the neural circuit by stress hormones, as differential endocrinological responses were not observed (Figure 1H). This suggests that more specific neural processing happens in the circuitry. In the circuitry downstream of the VNO, sensory signals can converge and diverge from the AOB before reaching the VMH. The axon terminals of hundreds of VNO neurons expressing the same receptor converge into a few specific glomerulus structures (Belluscio et al., 1999; Del Punta et al., 2002; Rodriguez et al., 1999), and, in some cases, neurons expressing receptors in the same subfamily send axons into the same glomeruli (Wagner et al., 2006). Individual AOB mitral cells extend multiple primary dendrites into multiple glomeruli (Imamura et al., 2020; Larriva-Sahd, 2008; Takami and Graziadei, 1991; Wagner et al., 2006; Yonekura and Yokoi, 2008). In turn, efferent axons of the mitral cells project to multiple brain substrates, namely the BNST, the MEA, and the posteromedial cortical amygdaloid nucleus (PMCo) (Davis et al., 1978; von Campenhausen and Mori, 2000). These structural features indicate that sensory inputs to the AOB mitral cells are converged, and outputs from the mitral cells may be diverged.

Earlier version: These features of the ascending neural pathway within the AOB may contribute to the sorting of sensory signals derived from fresh and old saliva. Fresh saliva activates more V2R-A4-expressing VNO neurons than old saliva, which may result in stronger activation of V2R-A4 glomeruli in fresh saliva-exposed mice compared to those exposed to old saliva. This differential activation of V2R-A4 glomeruli by fresh and old saliva may lead to activation of the mitral cells innervating the V2R-A4 glomeruli with different excitation levels. We assume that such differential activations of the mitral cells between fresh and old saliva result in the differential activation of targeting neural substrates, possibly MeApv, which results in differential activation of VMH neurons (Figure 7). The methodology that we undertook in this study, cFos-immunohistochemistry, only provides information about neurons responding to cat saliva above a given threshold. It is therefore important to examine how the mitral cells respond to fresh and old saliva using more quantitative techniques, such as in vivo electrophysiological recording, in future studies.

Revised version: These features of the ascending neural pathway within the AOB may contribute to the sorting of sensory signals derived from fresh and old saliva. Fresh saliva activates more V2R-A4-expressing VNO neurons than old saliva, which may result in stronger activation of V2R-A4 glomeruli in fresh saliva-exposed mice compared to those exposed to old saliva. We assume that such differential activations of V2R-A4 glomeruli between fresh and old saliva result in the differential activation of VMH neurons (Figure 8). The methodology that we undertook in this study, cFos-immunohistochemistry, only provides information about neurons responding to cat saliva above a given threshold. It is therefore important to examine how the neurons in this receptor circuitry respond to fresh and old saliva using more quantitative techniques, such as in vivo electrophysiological recording, in future studies.

Differential activation of VMH neurons potentially underlying distinct intensities of freezing behavior

Earlier version: Our findings indicate that the imminent predator cues present in cat saliva activate a specific subpopulation of VMH neurons, whose overall activity positively correlates with the intensity of freezing behavioral outputs (Figure 5E). In contrast, the less imminent predator cues activate a subpopulation of VMHdm/c neurons, whose overall activity does not correlate with the intensity of freezing behavior (Figure 5E). These findings suggest that induction of robust freezing behavior is determined by a distinct set of VMH neurons, and the intensity of freezing behavior is encoded separately, possibly by the intensity or dynamics of the neuronal firing, which cannot be detected by the cFos-IR used in this study. The VMHdm/c region contains several subpopulations of neurons that coordinate various defensive behavioral responses (Engelke et al., 2021; Esteban Masferrer et al., 2020; Kennedy et al., 2020; Silva et al., 2016; Tobias et al., 2023). Previous studies have demonstrated that optogenetic stimulation of VMHdm/c neurons induces freezing/immobility responses, suggesting that these neurons mediate freezing behavior in response to predator cues (Kunwar et al., 2015; Wang et al., 2015). In our current study, the distribution of cFos-IR in the VMH revealed a population of neurons clustered in close proximity that are more sensitive to fresh saliva than old saliva (Figure 6B). Therefore, we propose that imminent predator cues in cat saliva may activate these specific neurons with certain excitation levels to directly induce robust freezing behavior, while less imminent predator cues in cat saliva may use a different mechanism to regulate freezing behavior.

Revised version: Our findings indicate that the imminent predator cues present in cat saliva activate a specific subpopulation of VMH neurons, whose overall activity positively correlates with the intensity of freezing behavioral outputs (Figure 5E). In contrast, the less imminent predator cues activate a subpopulation of VMH neurons, whose overall activity does not correlate with the intensity of freezing behavior (Figure 5E). These findings suggest that induction of robust freezing behavior is determined by a distinct set of VMH neurons, and the intensity of freezing behavior is encoded separately, possibly by the intensity or dynamics of the neuronal firing, which cannot be detected by the cFos-IR used in this study. The VMHdm/c region contains several subpopulations of neurons that coordinate various defensive behavioral responses (Engelke et al., 2021; Esteban Masferrer et al., 2020; Kennedy et al., 2020; Silva et al., 2016; Tobias et al., 2023). Previous studies have demonstrated that optogenetic stimulation of VMHdm/c neurons induces freezing/immobility responses, suggesting that these neurons mediate freezing behavior in response to predator cues (Kunwar et al., 2015; Wang et al., 2015). In our current study, the distribution of cFos-IR in the VMH revealed a population of neurons clustered in close proximity that are more sensitive to fresh saliva than old saliva (Figure 6B). Therefore, we propose that imminent predator cues in cat saliva may activate these specific neurons with certain excitation levels to directly induce robust freezing behavior, while less imminent predator cues in cat saliva may use a different mechanism to regulate freezing behavior.

What distinguishes the subpopulation of VMHdm/c neurons activated by fresh saliva from those activated by old saliva? One possibility is that these neurons stimulate distinct downstream neural circuitries. The VMHdm/c neurons project to two major nuclei, the anterior hypothalamic nucleus (AHN) and the periaqueductal gray (PAG), which directly influence different aspects of defensive behavior (Canteras et al., 1994; Wang et al., 2015). Specifically, selective activation of the VMHdm/c-to-PAG pathway has been shown to cause freezing/immobility behavior, while activation of the VMHdm/c-to-AHN pathway rather promotes avoidance and induces immobility to a lesser extent (Wang et al., 2015). Although the number of cFos-positive cells in the dPAG did not seem to correlate with the intensity of freezing behavior in any saliva exposure group in our study (Figure 5F), it is possible that fresh saliva-activated neurons send inputs to the PAG, whereas old saliva-activated neurons send inputs to the AHN and/or other projection sites. Thus, the number of old saliva-activated cells in the VMH may correlate with different aspects of defensive responses, rather than freezing behavior.

Removed in the revised version: Given the observation from a previous study that neurons projecting to the PAG also send collateral projections to the AHN, and vice versa, neuronal populations activated by fresh and old saliva may or may not overlap. If the two sets of neurons overlap, it is possible that the strength of presynaptic inputs to the VMHdm/c differs between fresh and old saliva stimulation, which leads to differential activation of downstream pathways. Previous studies have shown that optogenetic stimulation of the VMHdm/c with different stimulus frequencies can induce different types of defensive behaviors (Wang et al., 2015), suggesting that distinct activity patterns of fresh and old saliva-responsive neurons in the VMHdm/c may contribute to differences in defensive responses. Investigating these activity patterns would be an intriguing direction for future research.

In summary, our study proposes a predator defensive circuitry that involves the accessory olfactory system and the VMHdm/c. Understanding the precise neural mechanisms of how different degrees of imminence of predator threat are processed in this circuit will provide valuable insights into the decision-making processes underlying defensive behaviors in animals.

Materials and methods

Animals

Earlier version: The experimental procedures were approved by the Institutional Animal Care and Use Committee at the University of California, Riverside (UCR) and were conducted in accordance with the National Institutes of Health Guide for the Care and Use of Laboratory Animals. C57BL/6N mice originally obtained from Charles River (Catalog# 027) were bred at the UCR vivarium. The Trpc2 knock-out mouse line was generously provided by Dr. Richard Axel Lab at Columbia University. All mice were housed in the vivarium under a reversed light cycle (lights on at 21:00; lights off at 9:00) for a minimum of 2 weeks prior to the behavior and histology experiments. Cats used in this study were obtained from the Nutrition and Pet Care Center at the University of California, Davis (a class A dealer) and were also housed in the UCR vivarium.

Revised version: The experimental procedures were approved by the Institutional Animal Care and Use Committee at the University of California, Riverside (UCR) and were conducted in accordance with the National Institutes of Health Guide for the Care and Use of Laboratory Animals. C57BL/6N mice originally obtained from Charles River (Catalog# 027) were bred at the UCR vivarium. The Trpc2 knock-out mouse line was generously provided by Dr. Richard Axel Lab at Columbia University. Fos2A-iCreERT2 (#030323) and Ai9 (#007909) mice were originally obtained from The Jackson Laboratory and were bred at the UCR vivarium. Fos2A-iCreERT2 mice were back-crossed with C57BL/6N strains of mice more than seven generations and bred with Ai9 mice to generate Fos2A-iCreERT2; Ai9. All mice were housed in the vivarium under a reversed light cycle (lights on at 21:00; lights off at 9:00) for a minimum of 2 weeks prior to the behavior and histology experiments. Cats used in this study were obtained from the Nutrition and Pet Care Center at the University of California, Davis (a class A dealer) and were also housed in the UCR vivarium.

Collection of saliva from cats

Earlier version: A person familiar with the cat gently held it for all procedures. A sterile cotton swab was inserted into the cat’s mouth where the cat chewed and licked on the swab for at least 20 seconds or until the swab is saturated with saliva. The swabs containing saliva were then placed in microcentrifuge tubes and stored until used. Fresh saliva was collected within 15 minutes prior to its use in the mouse behavioral study, while old saliva was collected 4 hours prior to the study. The old saliva was kept at room temperature until it was used.

Revised version: A person familiar with the cat gently held it for all procedures. A sterile cotton swab was inserted into the cat’s mouth where the cat chewed and licked on the swab for at least 20 seconds or until the swab is saturated with saliva. The swabs containing saliva were then placed in tightly sealed microcentrifuge tubes and stored until used. Fresh saliva was collected within 15 minutes prior to its use in the mouse behavioral study, while old saliva was collected 4 hours prior to the study. The old saliva was kept at room temperature until it was used.

Added in the revised version: To measure the volume of saliva collected, a saliva-containing swab was placed onto a custom filter made from 0.5 mL and 1.5 mL microcentrifuge tubes and centrifuged at 2,000 g for 1 minute using a tabletop mini centrifuge. The protein concentrations of the eluted saliva were measured using the Bio-Rad Protein Assay (5000001, Bio-Rad).

Added in the revised version: Western blot analysis

Added in the revised version: Proteins in saliva (10 μg protein) were separated by SDS-PAGE gels (456-1043, Bio-Rad, Hercules, CA), transferred onto PVDF membranes (162-0174, Bio-Rad), and incubated with rabbit anti-Fel d 4 antibody (1/1000: PA-FD4, Inbio). Bound antibody was detected by horseradish peroxidase-conjugated anti-rabbit IgG antibody (1/5,000: 7074, Cell Signaling) and developed using Clarity Western ECL Substrate (1705060, Bio-Rad). Signals were detected and captured using ImageQuant LAS 4000 mini (GE Healthcare, Pittsburgh, PA), and band intensities were quantified with ImageJ software. Statistical significance was determined by unpaired t-test (two-tailed).

Behavioral assay

All behavioral experiments were conducted during the animals’ dark cycle. Mice were placed in their experimental room and allowed approximately 4 hours to acclimate to the ambient conditions. A cat saliva-containing swab was manually introduced into the animal’s home cage to assess the defensive behavioral responses of resident mice. Recording of mouse behaviors commenced 2 minutes prior to introducing the sample and concluded 10 minutes thereafter. If tissues were required for histology, the animals were euthanized after 60 minutes, while for plasma collection, euthanasia took place after 15 minutes. Following the introduction of the sample, the behavioral room remained vacant of any individuals until the animal was euthanized. Mouse behavior was captured by night vision cameras equipped with infrared lights.

Behavioral analysis

Recordings of each behavioral session were saved and provided in a blinded manner to a researcher for analysis. The analysis was performed using a behavioral analysis program called Behavioral Observation Research Interactive Software (BORIS) (Friard and Gamba, 2016). Freezing behavior was defined as immobility lasting more than 2 seconds and was scored accordingly. Stretch sniffing risk assessment behavior was scored when animals exhibited flat-back/stretch-attend posture. Direct investigation was scored when there was physical contact with the swab. Animals that exhibited inactivity in their home cage for more than 50% of the pre-drop analysis duration in each experiment were excluded from the data set.

Histology and Immunohistochemistry (IHC)

After being exposed to the cat specimen or control swab for 60 minutes, the animals were euthanized with an overdose of sodium pentobarbital. Intracardial perfusions were performed using 4% paraformaldehyde (PFA) in PBS. The brains were extracted and post-fixed in 4% PFA overnight at 4°C. Subsequently, the brains were cryoprotected by sequentially immersing them in 15% and 30% sucrose solutions in PBS. The VNOs were also extracted and underwent the same post-fixation procedure, followed by decalcification with EDTA pH 8.0. Cryoprotection of the VNOs was achieved using a 30% sucrose solution in PBS. The cryoprotected tissues were then embedded in OCT (Optimal Cutting Temperature) media and frozen over liquid nitrogen. Tissues embedded in OCT were sectioned at 14 μm (VNO) and 20 μm (brain) thicknesses using a cryostat set at -25°C. The sections were placed onto positively charged microscope slides. The slides were stored at -80°C until stained. A standard frozen IHC protocol was employed. For VNO sections, staining was performed using a rabbit anti-pS6 antibody (Cell Signaling Technology, #4858) at a dilution of 1:200, incubated overnight at 4°C, followed by a secondary antibody, Goat anti-rabbit IgG Alexa Fluor 647 (Life Technologies #a21245), for 60 minutes at room temperature. For AOB sections, staining was carried out using a mouse anti-Gαi2 antibody (Millipore #MAB3077MI) and a guinea pig anti-cFos antibody (Synaptic Systems #226 308) at dilutions of 1:500 and 1:1000, respectively. This was followed by incubation with a goat anti-mouse IgG2b Alexa Fluor 546 (Life Technologies #a21143) and a Goat anti-guinea pig IgG Alexa Fluor 647. For brain sections, staining was performed using a guinea pig anti-cFos antibody at dilutions of 1:1000 incubated overnight at 4°C, followed by Goat anti-guinea pig IgG Alexa Fluor 647. Finally, the slides were mounted using MS Shield Mounting Medium With 4,6-diamidino-2-phenylindole (DAPI) and DABCO™ (EMS #17989-20).

Enzyme-linked Immunosorbent Assay (ELISA)

Plasma collected from mice exposed to cat saliva or control swabs was frozen at -80°C before initiating the ELISA protocol. The ACTH ELISA kit (MD Bioproducts) was utilized, following the instructions provided with the kit.

In situ hybridization (ISH) and dual staining

Detection of GαO RNA in fixed frozen VNO sections was conducted using the ACD RNAscope® probe targeting GNAO1 (GαO) (ACD #444991) and the RNAscope® Multiplex Fluorescent Reagent Kit v2 (ACD #323100). Probe binding was visualized using Akoya Biosciences’ Opal 570 (FP1488001KT) Dye at a dilution of 1:750 in RNAscope TSA Buffer. Following RNAscope staining, the slides were incubated in anti-pS6 antibody at 1:200 dilution overnight at 4 C, then detected by goat anti-rabbit IgG AlexaFluor 647 for one hour at room temperature.

Vmn2r (V2R) probes for the following genes were used: Vmn2r13 for subfamily A1, Vmn2r118 for A2, Vmn2r117 for A3, Vmn2r42 for A4, Vmn2r69 for A5-1, Vmn2r65 for A5-2, Vmn2r76 for A5-3, Vmn2r120 for A6, Vmn2r98 for A8-1, Vmn2r105 for A8-2, Vmn2r58 for A8-3, Vmn2r62 for A8-4, Vmn2r81 for A9, Vmn2r18 for A10, Vmn2r24 for B, and Vmn2r53 for D. Each cRNA probe was synthesized using the PCR-based method (Hua et al., 2018).

For detection of V2R RNA in ISH, VNO tissues were sectioned and placed onto slides up to one week before ISH and stored at -80 C. Fluorescent ISH was performed following published protocols (Ishii et al., 2004). Briefly, the VNO slides were baked at 60 C for 5 minutes and fixed with 4% PFA. The slides were then treated with 10 µg/mL Proteinase K (Sigma #3115828001), followed by being acetylated in 0.1 M Triethanolamine (Sigma #90279)-HCl with acetic anhydride (Sigma #320102). V2R cRNA probes were hybridized with tissue sections overnight at 65 C in a hybridization solution consisting of 50 % formamide, 10 mM Tris-HCl pH 8.0, 200 µg/mL yeast tRNA, 10 % dextran sulfate, 1x Denhardt’s, 600 mM NaCl, 0.25 % SDS, 1 mM EDTA pH 8.0. After hybridization, the slides were washed with 2x SSC in 50 % formamide for 30 minutes, then 2x SSC, 0.2x SSC, and 0.1x SSC in H2O for 20 minutes each at 65 C. The slides were then blocked with blocking reagent (Perkin Elmer) and 0.5 % peroxidase block sequentially. Anti-DIG-POD (Millipore Sigma #11207733910) was used to detect DIG-labeled probes. TSA-Cy3 Plus kit (Akoya #NEL744001KT) was used to detect anti-DIG antibody at 1:200. The slides were incubated with anti-pS6 antibody at 1:200 dilution overnight at 4 C, then detected by goat anti-rabbit IgG AlexaFluor 647 for one hour at room temperature.

Following staining, the slides were mounted with MS Shield Mounting Medium with DAPI and DABCO™.

Image Quantification

Fluorescent images were captured at either 10X or 20X magnification using the Zeiss Axio Imager.M2 microscope. Whole slide images of IHC-stained brains were scanned through fluorescent imaging using Pannoramic SCAN (3D Histech) at Reveal Biosciences (San Diego, CA). Automated counting of fluorescent positive cells was performed using a proprietary code developed in our laboratory within the QuPath software.

VMH cFos IR Quantification and Distribution Analysis

All VMH regions were identified by clusters of nuclei visualized through DAPI staining and manually annotated in QuPath using standard coronal mouse brain atlas as reference. Annotations were processed with a proprietary script for immunoreactivity quantification. Distribution analyses were performed in QuPath using a rectangular annotation set parallel to the dorsomedial-ventrolateral axis of each VMH region. This rectangular annotation was then split into 10 equal bins, and immunoreactivity quantification in each bin was automated in QuPath, where the detections were finally manually recorded.

Added in the revised version: Double exposure experiment

Added in the revised version: All behavioral experiments were conducted during the animals’ dark cycle. Individual Fos2A-iCreERT2; Ai9 mice in their home cages were placed in an experimental room and allowed to acclimate to the ambient conditions for approximately 4 hours. The 4-Hydroxytamoxifen (4-OHT) solution was prepared as previously described (DeNardo et al., 2019). During the first exposure session, mice were injected intraperitoneally with 50 mg/kg of 4-OHT, 30 minutes prior to the introduction of the saliva sample. The first swab was removed from the cage 10 minutes after a mouse made the initial direct interaction with it. The mouse then remained in the cage for an additional 90 minutes before being returned to the vivarium. The second exposure session was conducted one week after the first session. For the second exposure, the mouse was kept with a swab for 60 minutes before being prepared for histological analysis. The behavioral room remained unoccupied during all exposure periods.

Statistical Analysis

Earlier version: All statistical analyses were performed using Prism software (GraphPad). Behavioral analysis results and cFos-IR distribution in the VMH, involving two variables, were compared using a traditional Two-Way ANOVA with Tukey’s multiple comparisons. Behavioral and immunohistological results were statistically compared using One-Way ANOVA with Tukey’s multiple comparisons, assuming a parametric normal distribution. The correlation between immunoreactivity and behavior was analyzed using either Spearman’s rank correlation coefficient or Pearson correlation coefficient: for data points that did not pass Shapiro-Wilk and Kolmogorov-Smirnov normality tests, Spearman’s rank correlation coefficient was used, while for data points that pass Shapiro-Wilk and Kolmogorov-Smirnov normality tests, Pearson correlation coefficient was used.

Revised version: Statistical analyses were performed using Prism software (GraphPad) or MATLAB (MathWorks). Behavioral analysis results and cFos-IR distribution in the VMH, involving two variables, were compared using a traditional Two-Way ANOVA with Tukey’s multiple comparisons. Behavioral and immunohistological results were statistically compared using One-Way ANOVA with Tukey’s multiple comparisons, assuming a parametric normal distribution. The correlation between immunoreactivity and behavior was analyzed using either Spearman’s rank correlation coefficient or Pearson correlation coefficient: for data points that did not pass Shapiro-Wilk and Kolmogorov-Smirnov normality tests, Spearman’s rank correlation coefficient was used, while for data points that pass Shapiro-Wilk and Kolmogorov-Smirnov normality tests, Pearson correlation coefficient was used. Statistical comparisons of activated neural populations in the double exposure experiment were conducted using permutation tests and Chi-squared tests.

Removed in the revised version: Supporting information